[1506392015] SERVICE ALERT: dev-nginx;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506392075] SERVICE ALERT: dev-nginx;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506392135] SERVICE ALERT: dev-nginx;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506392195] SERVICE ALERT: dev-nginx;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416202] SERVICE ALERT: dev-bastion;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.001 second response time
[1506416210] SERVICE ALERT: dev-database;HTTP;CRITICAL;SOFT;1;connect to address 10.1.101.10 and port 80: Connection refused
[1506416225] SERVICE ALERT: dev-services;HTTP;CRITICAL;SOFT;1;connect to address 10.1.100.10 and port 80: Connection refused
[1506416262] SERVICE ALERT: dev-bastion;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.001 second response time
[1506416270] SERVICE ALERT: dev-database;HTTP;CRITICAL;SOFT;2;connect to address 10.1.101.10 and port 80: Connection refused
[1506416285] SERVICE ALERT: dev-services;HTTP;CRITICAL;SOFT;2;connect to address 10.1.100.10 and port 80: Connection refused
[1506416315] SERVICE ALERT: dev-bastion;SSH;CRITICAL;SOFT;1;connect to address 10.1.0.10 and port 22: Connection refused
[1506416322] SERVICE ALERT: dev-bastion;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.000 second response time
[1506416330] SERVICE ALERT: dev-database;HTTP;CRITICAL;SOFT;3;connect to address 10.1.101.10 and port 80: Connection refused
[1506416345] SERVICE ALERT: dev-services;HTTP;CRITICAL;SOFT;3;connect to address 10.1.100.10 and port 80: Connection refused
[1506416352] SERVICE ALERT: dev-bastion;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416360] SERVICE ALERT: dev-database;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416375] SERVICE ALERT: dev-services;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416375] SERVICE ALERT: dev-bastion;SSH;CRITICAL;SOFT;2;connect to address 10.1.0.10 and port 22: Connection refused
[1506416382] SERVICE ALERT: dev-bastion;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 403 Forbidden - 5179 bytes in 0.000 second response time
[1506416390] SERVICE ALERT: dev-database;HTTP;CRITICAL;HARD;4;connect to address 10.1.101.10 and port 80: Connection refused
[1506416405] SERVICE ALERT: dev-services;HTTP;CRITICAL;HARD;4;connect to address 10.1.100.10 and port 80: Connection refused
[1506416412] SERVICE ALERT: dev-bastion;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416420] SERVICE ALERT: dev-database;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416435] SERVICE ALERT: dev-services;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416435] SERVICE ALERT: dev-bastion;SSH;CRITICAL;SOFT;3;connect to address 10.1.0.10 and port 22: Connection refused
[1506416472] SERVICE ALERT: dev-bastion;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416480] SERVICE ALERT: dev-database;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416495] SERVICE ALERT: dev-services;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416495] SERVICE ALERT: dev-bastion;SSH;CRITICAL;HARD;4;connect to address 10.1.0.10 and port 22: Connection refused
[1506416532] SERVICE ALERT: dev-bastion;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416540] SERVICE ALERT: dev-database;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506416555] SERVICE ALERT: dev-services;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1506438395] SERVICE ALERT: dev-bastion;SSH;OK;HARD;4;SSH OK - OpenSSH_6.6.1 (protocol 2.0)